Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Collection Algorithms Advanced

In this part your are going to work with advanced concepts on collections.

## Exercise (2 Points)

In this exercise we practice working with matrices, i.e. lists of lists.

### Task 1

Given two integer parameters `n` and `m`, write a function that returns a matrix `table` with `n` columns and `m` rows filled with numbers between `1` and `n*m` in the order left-to right, top to bottom. For example, for `n = 5` and `m = 6` the matrix should look as follows:
```
[[ 1,  2,  3,  4,  5,  6], 
[ 7,  8,  9, 10, 11, 12], 
[13, 14, 15, 16, 17, 18], 
[19, 20, 21, 22, 23, 24], 
[25, 26, 27, 28, 29, 30]]
```

<!-- grade_id: cell-1e49c16c9b32e1d2, points: 0, grade: false, solution: true -->
```python
def create_matrix(n,m):
    table = []
# YOUR SOLUTION HERE
    return table
```

<!-- grade_id: cell-399e01d64b13268c, points: 1, grade: true, solution: false -->
```python
# This tests your implementation:
assert create_matrix(5,6) == [
     [1, 2, 3, 4, 5, 6],
     [7, 8, 9, 10, 11, 12],
     [13, 14, 15, 16, 17, 18],
     [19, 20, 21, 22, 23, 24],
     [25, 26, 27, 28, 29, 30]
]
print('Success!')
```

### Task 2

Write a function that returns an (ordinary 1-dimensional) list `columnProducts` that contains the product of elements of the matrix `table` for each column of this matrix. For the example matrix given in Task 1 the array should be:

```
[43225, 116480, 229635, 394240, 623645, 933120]
```

<!-- grade_id: cell-ba9b91a79aeb711d, points: 0, grade: false, solution: true -->
```python
def column_products(table):
    column_products = []
# YOUR SOLUTION HERE
    return column_products
```

<!-- grade_id: cell-eade967df283b142, points: 1, grade: true, solution: false -->
```python
# This tests your implementation:
table = [
     [1, 2, 3, 4, 5, 6],
     [7, 8, 9, 10, 11, 12],
     [13, 14, 15, 16, 17, 18],
     [19, 20, 21, 22, 23, 24],
     [25, 26, 27, 28, 29, 30]
]
assert column_products(table) == [43225, 116480, 229635, 394240, 623645, 933120]
print('Success!')
```

## Exercise (2 Points)

In this exercise we deal with the problems arising from references and variable bindings.
In lecture 8 we discussed the difference between assignment of list variables and assignment of list values. In particular, since lists are internally represented as pointers to memory locations, it may happen that changing a value of one list changes the value of another list. In this exercise we practice with understanding and predicting such effects.

Consider the following Python program that operates with lists:

```python
a = [0,1]
b = a[1]
b += 1
a[0] = b
c = [1,0]
for i in range(2): c[i] = a[i]
a[0] = a[1]
a[1] += 2
d = a
d[0] += d[1]
a[0] = d[1]
d[1] += 2
a = c
c = None
```

Simulate the execution of this program by completing the following table with the values of each variable after execution of the statement on the corresponding line. (Line numbers can be enabled from the menu `View` $\to$ `Show Line Numbers`)

Write `?` for a variable that has not been initialized yet, `None` for variables with the value `None` and provide array contents as comma separated lists in square brackets. The first two lines are already given as examples. Double click on this cell, copy the table to the answer cell below, and add extend the table with your answers.

| Line | a   | b | c | d |
|:----:|:---:|:-:|:-:|:-:|
|1     |[0,1]| ? | ? | ? |
|2     |[0,1]| 1 | ? | ? |

<!-- grade_id: cell-238eb3de3d3fda9b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE