Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Text Processing

Many text editors can perform basic manipulations with text, such as searching and/or replacing characters or substrings. Recall how we have edited a Python program in the [part 3 of Assignment #2](../assignment02/part3.ipynb).
In this part we perform some simple operations with text using Python conditions.

In some of the tasks below, you are asked to write a function that returns the given text where specified characters were replaced with the underscore `_`. Usually you are asked to replace all *vowels*, i.e., letters: `a`, `e`, `i`, `o`, `u`, in the given text. For example, for the text

```
Hello World!
```
your function should return:

```
H_ll_ W_rld!
```

Some parts of the functions are already provided for your convenience. **Please do not remove or modify these parts!** Please carefully read the instruction for each task.

**Important: As always you are only allowed to use features that were covered up to this point in the lecture, i.e. are part of lecture 7! This means you are not allowed to use build-in functions that operate on strings, execpt for `len()`, `ord()`, `chr()` and `<str>.format()` (as well as f-strings).**

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each vowel with a dash (`-`) using an `if-else` statement with a **single condition**, that consists of a single `in` operator.

<!-- grade_id: cell-38241276a7919b5a, points: 0, grade: false, solution: true -->
```python
def vowels(text):
    new_string = ""
    for c in text:
      if c in "AEIOUaeiou":  
            new_string += "-"  
      else:
        new_string += c    

    return new_string
```

```python
def vowels(text): #Text = Hello World
    new_string = ""
    vowels = "aeiouAEIOU"
    for x in text:
        if x in vowels:
            new_string += "_"
        else:
            new_string += x
    return new_string
    
#string = "Hello World"
#vowels(string)

```

<!-- grade_id: cell-f2db875188f9b58e, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(vowels("Hello World"))) == """<class 'str'>""", """type of vowels(\"Hello World\") is not str. vowels(\"Hello World\") should be an str"""
assert str(len(vowels("Hello World"))) == """11""", """length of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World").lower()) == """h-ll- w-rld""", """value of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World")) == """H-ll- W-rld""", """correct string value of vowels(\"Hello World\") but incorrect case of letters"""

print('Success!')
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[21], line 5
      3 assert str(type(vowels("Hello World"))) == """<class 'str'>""", """type of vowels(\"Hello World\") is not str. vowels(\"Hello World\") should be an str"""
      4 assert str(len(vowels("Hello World"))) == """11""", """length of vowels(\"Hello World\") is not correct"""
----> 5 assert str(vowels("Hello World").lower()) == """h-ll- w-rld""", """value of vowels(\"Hello World\") is not correct"""
      6 assert str(vowels("Hello World")) == """H-ll- W-rld""", """correct string value of vowels(\"Hello World\") but incorrect case of letters"""
      8 print('Success!')

AssertionError: value of vowels("Hello World") is not correct
```

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each **letter appearing in the given string** `replace` with a dash. Note: you should replace the text given in the variable `text`, *not* the text in `replace`.


*Warning:* Your program should work **regardless** which letters are given in the `replace` string, and even if no letters are provided at all (when `replace` is the empty string).

<!-- grade_id: cell-dba35b06c5e38d24, points: 0, grade: false, solution: true -->
```python
def general_replace(text, replace):
    new_string = ""
    for c in text:
        if c in replace:
            new_string += "-"
        else:
            new_string += c
    return new_string

```

<!-- grade_id: cell-ea7b28540a173711, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(general_replace("Hello World", "aeiou"))) == """<class 'str'>""", """type of general_replace(\"Hello World\", \"aeiou\") is not str. general_replace(\"Hello World\", \"aeiou\") should be an str"""
assert str(len(general_replace("Hello World", "aeiou"))) == """11""", """length of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou").lower()) == """h-ll- w-rld""", """value of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou")) == """H-ll- W-rld""", """correct string value of general_replace(\"Hello World\", \"aeiou\") but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (1P.)

For the given string `text` return a string that consists of `n`-times the original string, for example for "a" and 5 you should return "aaaaa".

<!-- grade_id: cell-27137fae67d5c4b5, points: 0, grade: false, solution: true -->
```python
def repeat(text, n):
    return (text*n)
```

<!-- grade_id: cell-8defb15d0b45d3f9, points: 1, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(repeat("a", 5))) == """<class 'str'>""", """type of repeat(\"a\", 5) is not str. repeat(\"a\", 5) should be an str"""
assert str(len(repeat("a", 5))) == """5""", """length of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5).lower()) == """aaaaa""", """value of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5)) == """aaaaa""", """correct string value of repeat(\"a\", 5) but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

For the given input `text`, check whether the string in `text` is the same as its reverse (the string read backwards). For example, your function should return `True` for "level" but not for "lever".

*Hint: You can use slice notation for the reversing.*

```python
text = input("What is the word to reverse?")
print(text[::-1])
```

Output:
```
What is the word to reverse? ty
yt
```

<!-- grade_id: cell-d848906595b1a870, points: 0, grade: false, solution: true -->
```python
def backwards_equality(text):
    reverse = text[::-1]
    if reverse == text:
        return True
    else:
        return False
        
backwards_equality("lever")

```

Output:
```
False
```

<!-- grade_id: cell-20b294b518a836b2, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(backwards_equality("level"))) == """<class 'bool'>""", """type of backwards_equality(\"level\") is not bool. backwards_equality(\"level\") should be a bool"""
assert str(backwards_equality("level")) == """True""", """boolean value of backwards_equality(\"level\") is not correct"""

assert str(type(backwards_equality("lever"))) == """<class 'bool'>""", """type of backwards_equality(\"lever\") is not bool. backwards_equality(\"lever\") should be a bool"""
assert str(backwards_equality("lever")) == """False""", """boolean value of backwards_equality(\"lever\") is not correct"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

In the return variable `new_string`, return a string where the substring of the given string `s` between the indices `start` and `end` (inclusive) is reversed, while the rest of the string remains unchanged.  

**Make sure that your function returns a reasonable result for every input (assume `start` and `end` are valid indices, and `start <= end`).**

*Hint: Use string slicing and indexing.*

<!-- grade_id: cell-70c2a1b68e7165d8, points: 0, grade: false, solution: true -->
```python
def reverse_substring(s, start, end):
    substring = s[start : end + 1]
    new_string = substring[::-1]
    return new_string

reverse_substring("Hello World!", 3, 6)
```

Output:
```
'W ol'
```

<!-- grade_id: cell-c793f233248fc05c, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(reverse_substring("Hello World!", 3, 6))) == """<class 'str'>""", """type of reverse_substring(\"Hello World!\", 3, 6) is not str. reverse_substring(\"Hello World!\", 3, 6) should be an str"""
assert str(len(reverse_substring("Hello World!", 3, 6))) == """12""", """length of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6).lower()) == """helw olorld!""", """value of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6)) == """HelW olorld!""", """correct string value of reverse_substring(\"Hello World!\", 3, 6) but incorrect case of letters"""

assert str(type(reverse_substring("Hi", 0, 1))) == """<class 'str'>""", """type of reverse_substring(\"Hi\", 0, 1) is not str. reverse_substring(\"Hi\", 0, 1) should be an str"""
assert str(len(reverse_substring("Hi", 0, 1))) == """2""", """length of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1).lower()) == """ih""", """value of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1)) == """iH""", """correct string value of reverse_substring(\"Hi\", 0, 1) but incorrect case of letters"""

print('Success!')
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[31], line 4
      2 from hashlib import sha1
      3 assert str(type(reverse_substring("Hello World!", 3, 6))) == """<class 'str'>""", """type of reverse_substring(\"Hello World!\", 3, 6) is not str. reverse_substring(\"Hello World!\", 3, 6) should be an str"""
----> 4 assert str(len(reverse_substring("Hello World!", 3, 6))) == """12""", """length of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
      5 assert str(reverse_substring("Hello World!", 3, 6).lower()) == """helw olorld!""", """value of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
      6 assert str(reverse_substring("Hello World!", 3, 6)) == """HelW olorld!""", """correct string value of reverse_substring(\"Hello World!\", 3, 6) but incorrect case of letters"""

AssertionError: length of reverse_substring("Hello World!", 3, 6) is not correct
```

## Task (2P.)

Write a function, that given two strings `text1` and `text2`, checks whether they are equal if case is ignored, so "hELLO" should be equal to "HellO".

*Remember that you are not allowed to use build-in string methods such as `s.isupper()`.*

*Hint: Use lexicographic comparison.*

*Hint: Remember `break` and `continue`.*

<!-- grade_id: cell-8a5009af8b821f0f, points: 0, grade: false, solution: true -->
```python
def equal_without_case(text1, text2):
    if len(text1) != len(text2):
        return False
    for c1, c2 in zip(text1, text2):
        if c1.lower() != c2.lower():
            return False
    return True

```

<!-- grade_id: cell-81dbc76ec9ba6ce1, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(equal_without_case("hELLO", "HellO"))) == """<class 'bool'>""", """type of equal_without_case(\"hELLO\", \"HellO\") is not bool. equal_without_case(\"hELLO\", \"HellO\") should be a bool"""
assert str(equal_without_case("hELLO", "HellO")) == """True""", """boolean value of equal_without_case(\"hELLO\", \"HellO\") is not correct"""

assert str(type(equal_without_case("hello", "HELLO"))) == """<class 'bool'>""", """type of equal_without_case(\"hello\", \"HELLO\") is not bool. equal_without_case(\"hello\", \"HELLO\") should be a bool"""
assert str(equal_without_case("hello", "HELLO")) == """True""", """boolean value of equal_without_case(\"hello\", \"HELLO\") is not correct"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

In this task, you are meant to re-create simple [ASCII art](https://en.wikipedia.org/wiki/ASCII_art).  
Write a function that returns the following pattern for different given heights:

### Example Height = 10

```
*        *
 *      *
  *    *
   *  *
++++++++++
   *  *
  *    *
 *      *
*        *
```

### Example Height = 7

```
*     *
 *   *
  * *
+++++++
  * *
 *   *
*     *
```

The function should receive the height as an input and return a string containing the pattern.

**Hints**:
1. To add a new line to a string `s`, use `s += "\n"`. To see the resulting string formatted correctly, you must use `print(s)`; evaluating `s` alone won’t display the line breaks correctly in JupyterHub.
2. Treat even and odd heights differently to maintain symmetry, as seen in the examples.
3. The lowest height for which the function must produce a meaningful string is `3`.

<!-- grade_id: cell-2d0414bef2c32926, points: 0, grade: false, solution: true -->
```python
def generate_pattern(height):
    pattern = ""
# YOUR SOLUTION HERE
    return pattern
```

<!-- grade_id: cell-1a8eb14a86a88b76, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(generate_pattern(10))) == """<class 'str'>""", """type of generate_pattern(10) is not str. generate_pattern(10) should be an str"""
assert str(len(generate_pattern(10))) == """87""", """length of generate_pattern(10) is not correct"""
assert str(generate_pattern(10).lower()) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """value of generate_pattern(10) is not correct"""
assert str(generate_pattern(10)) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """correct string value of generate_pattern(10) but incorrect case of letters"""

assert str(type(generate_pattern(6))) == """<class 'str'>""", """type of generate_pattern(6) is not str. generate_pattern(6) should be an str"""
assert str(len(generate_pattern(6))) == """33""", """length of generate_pattern(6) is not correct"""
assert str(generate_pattern(6).lower()) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """value of generate_pattern(6) is not correct"""
assert str(generate_pattern(6)) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """correct string value of generate_pattern(6) but incorrect case of letters"""

assert str(type(generate_pattern(3))) == """<class 'str'>""", """type of generate_pattern(3) is not str. generate_pattern(3) should be an str"""
assert str(len(generate_pattern(3))) == """12""", """length of generate_pattern(3) is not correct"""
assert str(generate_pattern(3).lower()) == """* *\n+++\n* *\n""", """value of generate_pattern(3) is not correct"""
assert str(generate_pattern(3)) == """* *\n+++\n* *\n""", """correct string value of generate_pattern(3) but incorrect case of letters"""

print('Success!')
```