## **Python String Exercises (Beginner to Advanced)**

### **Beginner Level: Basic String Manipulation**



#### **1. Concatenation**
Write a program that takes two strings as input and prints them concatenated together.



```python
def con(s,m):
    return s + m

print(con("He","llo"))
```

Output:
```
Hello
```

#### **2. String Indexing & Slicing**
- Print the first and last character of a given string.
- Print a substring containing only the first 5 characters.



```python

```

```python

```

```python
string = "Hakehdy"
print(string[:len(string)])
```

Output:
```
Hakehdy
```

```python
string = "Hakehdy"
print(string[0],string[-1])
```

Output:
```
H y
```

```python
def first_and_last(string):
    for i in range(len(string)):
        if i == 0 or  i == len(string) -1:
            print(string[i], end = "")
    #return None

first_and_last("Hello")
```

Output:
```
Ho
```

```python
def prnt(st):
    for i in range(len(st)-1):
        print(st[i])
        
        
prnt("hello")
```

Output:
```
h
e
l
l
```

#### **3. String Methods**
Convert a given string to:
- Uppercase
- Lowercase
- Title case (first letter capitalized)



```python
# Using a method

def upperC(string):
    string = string.upper()
    return string

# Using a method

def upperC(string):
    string = string.upper()
    return string

print(upperC(string))

print(upperC(string))
```

Output:
```
HAKEHDY
HAKEHDY
```

```python
# Lexicographic with numbers
def up(string):
    new = ""
    for i in string:
        if 97< ord(i) <122 or 64 < ord(i) < 91:
            if 97 <= ord(i) <= 122:
                new += chr(ord(i)-32)
            else:
                new += i
        else:
            return ("Special letters are present")
    return new

print(up("hElLo"))
```

Output:
```
HELLO
```

```python

```

```python
# Lexicographic with alphabets
def up(string):
    new = ""
    for i in string:
        if "a" <= i <= "z":
            new += chr(ord(i)-32)
        else:
            new += i
    return new

print(up("hEllo"))
```

Output:
```
HELLO
```

```python
# Lexicographic lower with numbers
def low(string):
    new = ""
    for i in string:
        if 65 <= ord(i) <= 91:
            new += chr(ord(i)+32)
        else:
            new += i
    return new

print(low("heLLo"))
```

Output:
```
hello
```

```python
# Lexicographic lower with alpha
def low(string):
    new = ""
    for i in string:
        if "A" <= i <= "Z":
            new += chr(ord(i)+32)
        else:
            new += i
    return new

print(low("heLLo"))
```

Output:
```
hello
```

```python
# Title case
def titl(string):
    new = ""
    for i in range(len(string)):
        if i == 0:
            if 97 <= ord(string[i]) <= 122:
                new += chr(ord(string[i])-32)
            else:
                new += string[i]
        else:
            new += string[i]
    return new

print(titl("he"))
```

Output:
```
He
```

#### **4. Find and Replace**
- Replace all occurrences of the word **"apple"** with **"orange"** in a given string.

---



```python
def rep():
    
```

```python

```

### **Intermediate Level: Loops with Strings**

#### **5. Iterate Over a String**
Write a program that prints each character of a string on a new line.



```python

```

```python

```

#### **6. Count Vowels in a String**
Write a function that counts the number of vowels (`a, e, i, o, u`) in a string.



```python

```

```python

```

#### **7. Reverse a String Using a Loop**
Write a program to reverse a string manually using a loop.



```python

```

```python

```

#### **8. Check if a String is a Palindrome**
Write a function to check if a given string is a palindrome (same forward and backward).



```python

```

```python

```

#### **9. Find the Longest Word in a Sentence**
Write a function that takes a sentence as input and returns the longest word.

---



```python

```

```python

```

### **Advanced Level: Control Flow with Strings**

#### **10. Word Frequency Counter**
Count the occurrences of each word in a sentence and store the result in a dictionary.



```python

```

```python

```

#### **11. Extract Digits from a String**
Write a function that extracts only the digits from a given string.



```python

```

```python

```

#### **12. Remove Special Characters from a String**
Write a function that removes special characters (e.g., `@, !, #, $, %`) from a string.



```python

```

```python

```

#### **13. Check for Anagrams**
Write a function that checks if two strings are anagrams (contain the same characters in a different order).



```python

```

```python

```

#### **14. Capitalize Every Alternate Letter in a String**
Write a function that capitalizes every alternate letter in a given string.



```python

```

```python

```

#### **15. Find the Most Common Character in a String**
Write a function that finds and returns the most frequently occurring character in a string.

---


```python

```

```python

```

#### **16. Slicing**

Copy the provided `text` to the returned variable `new_string` after replacing each vowel with a dash (`-`) using an `if-else` statement with a **single condition**, that consists of a single `in` operator.

```python

```

```python

```

```python

```

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each vowel with a dash (`-`) using an `if-else` statement with a **single condition**, that consists of a single `in` operator.

<!-- grade_id: cell-38241276a7919b5a, points: 0, grade: false, solution: true -->
```python
def vowels(text):
  

    
```

Output:
```
Cell In[10], line 4
    
    ^
SyntaxError: incomplete input
```

<!-- grade_id: cell-f2db875188f9b58e, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(vowels("Hello World"))) == """<class 'str'>""", """type of vowels(\"Hello World\") is not str. vowels(\"Hello World\") should be an str"""
assert str(len(vowels("Hello World"))) == """11""", """length of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World").lower()) == """h-ll- w-rld""", """value of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World")) == """H-ll- W-rld""", """correct string value of vowels(\"Hello World\") but incorrect case of letters"""

print('Success!')
```

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each **letter appearing in the given string** `replace` with a dash. Note: you should replace the text given in the variable `text`, *not* the text in `replace`.


*Warning:* Your program should work **regardless** which letters are given in the `replace` string, and even if no letters are provided at all (when `replace` is the empty string).

<!-- grade_id: cell-dba35b06c5e38d24, points: 0, grade: false, solution: true -->
```python
def general_replace(text, replace):
   
```

<!-- grade_id: cell-ea7b28540a173711, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(general_replace("Hello World", "aeiou"))) == """<class 'str'>""", """type of general_replace(\"Hello World\", \"aeiou\") is not str. general_replace(\"Hello World\", \"aeiou\") should be an str"""
assert str(len(general_replace("Hello World", "aeiou"))) == """11""", """length of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou").lower()) == """h-ll- w-rld""", """value of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou")) == """H-ll- W-rld""", """correct string value of general_replace(\"Hello World\", \"aeiou\") but incorrect case of letters"""

print('Success!')
```

## Task (1P.)

For the given string `text` return a string that consists of `n`-times the original string, for example for "a" and 5 you should return "aaaaa".

<!-- grade_id: cell-27137fae67d5c4b5, points: 0, grade: false, solution: true -->
```python
def repeat(text, n):
    
```

<!-- grade_id: cell-8defb15d0b45d3f9, points: 1, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(repeat("a", 5))) == """<class 'str'>""", """type of repeat(\"a\", 5) is not str. repeat(\"a\", 5) should be an str"""
assert str(len(repeat("a", 5))) == """5""", """length of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5).lower()) == """aaaaa""", """value of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5)) == """aaaaa""", """correct string value of repeat(\"a\", 5) but incorrect case of letters"""

print('Success!')
```

## Task (2P.)

For the given input `text`, check whether the string in `text` is the same as its reverse (the string read backwards). For example, your function should return `True` for "level" but not for "lever".

*Hint: You can use slice notation for the reversing.*

<!-- grade_id: cell-d848906595b1a870, points: 0, grade: false, solution: true -->
```python
def backwards_equality(text):
   
```

<!-- grade_id: cell-20b294b518a836b2, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(backwards_equality("level"))) == """<class 'bool'>""", """type of backwards_equality(\"level\") is not bool. backwards_equality(\"level\") should be a bool"""
assert str(backwards_equality("level")) == """True""", """boolean value of backwards_equality(\"level\") is not correct"""

assert str(type(backwards_equality("lever"))) == """<class 'bool'>""", """type of backwards_equality(\"lever\") is not bool. backwards_equality(\"lever\") should be a bool"""
assert str(backwards_equality("lever")) == """False""", """boolean value of backwards_equality(\"lever\") is not correct"""

print('Success!')
```

## Task (2P.)

In the return variable `new_string`, return a string where the substring of the given string `s` between the indices `start` and `end` (inclusive) is reversed, while the rest of the string remains unchanged.  

**Make sure that your function returns a reasonable result for every input (assume `start` and `end` are valid indices, and `start <= end`).**

*Hint: Use string slicing and indexing.*

<!-- grade_id: cell-70c2a1b68e7165d8, points: 0, grade: false, solution: true -->
```python
def reverse_substring(s, start, end):
   
```

<!-- grade_id: cell-c793f233248fc05c, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(reverse_substring("Hello World!", 3, 6))) == """<class 'str'>""", """type of reverse_substring(\"Hello World!\", 3, 6) is not str. reverse_substring(\"Hello World!\", 3, 6) should be an str"""
assert str(len(reverse_substring("Hello World!", 3, 6))) == """12""", """length of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6).lower()) == """helw olorld!""", """value of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6)) == """HelW olorld!""", """correct string value of reverse_substring(\"Hello World!\", 3, 6) but incorrect case of letters"""

assert str(type(reverse_substring("Hi", 0, 1))) == """<class 'str'>""", """type of reverse_substring(\"Hi\", 0, 1) is not str. reverse_substring(\"Hi\", 0, 1) should be an str"""
assert str(len(reverse_substring("Hi", 0, 1))) == """2""", """length of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1).lower()) == """ih""", """value of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1)) == """iH""", """correct string value of reverse_substring(\"Hi\", 0, 1) but incorrect case of letters"""

print('Success!')
```

## Task (2P.)

Write a function, that given two strings `text1` and `text2`, checks whether they are equal if case is ignored, so "hELLO" should be equal to "HellO".

*Remember that you are not allowed to use build-in string methods such as `s.isupper()`.*

*Hint: Use lexicographic comparison.*

*Hint: Remember `break` and `continue`.*

<!-- grade_id: cell-8a5009af8b821f0f, points: 0, grade: false, solution: true -->
```python
def equal_without_case(text1, text2):
  

```

<!-- grade_id: cell-81dbc76ec9ba6ce1, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(equal_without_case("hELLO", "HellO"))) == """<class 'bool'>""", """type of equal_without_case(\"hELLO\", \"HellO\") is not bool. equal_without_case(\"hELLO\", \"HellO\") should be a bool"""
assert str(equal_without_case("hELLO", "HellO")) == """True""", """boolean value of equal_without_case(\"hELLO\", \"HellO\") is not correct"""

assert str(type(equal_without_case("hello", "HELLO"))) == """<class 'bool'>""", """type of equal_without_case(\"hello\", \"HELLO\") is not bool. equal_without_case(\"hello\", \"HELLO\") should be a bool"""
assert str(equal_without_case("hello", "HELLO")) == """True""", """boolean value of equal_without_case(\"hello\", \"HELLO\") is not correct"""

print('Success!')
```

## Task (2P.)

In this task, you are meant to re-create simple [ASCII art](https://en.wikipedia.org/wiki/ASCII_art).  
Write a function that returns the following pattern for different given heights:

### Example Height = 10

```
*        *
 *      *
  *    *
   *  *
++++++++++
   *  *
  *    *
 *      *
*        *
```

### Example Height = 7

```
*     *
 *   *
  * *
+++++++
  * *
 *   *
*     *
```

The function should receive the height as an input and return a string containing the pattern.

**Hints**:
1. To add a new line to a string `s`, use `s += "\n"`. To see the resulting string formatted correctly, you must use `print(s)`; evaluating `s` alone won’t display the line breaks correctly in JupyterHub.
2. Treat even and odd heights differently to maintain symmetry, as seen in the examples.
3. The lowest height for which the function must produce a meaningful string is `3`.

<!-- grade_id: cell-2d0414bef2c32926, points: 0, grade: false, solution: true -->
```python
def generate_pattern(height):
    pattern = ""
# YOUR SOLUTION HERE
    return pattern
```

<!-- grade_id: cell-1a8eb14a86a88b76, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(generate_pattern(10))) == """<class 'str'>""", """type of generate_pattern(10) is not str. generate_pattern(10) should be an str"""
assert str(len(generate_pattern(10))) == """87""", """length of generate_pattern(10) is not correct"""
assert str(generate_pattern(10).lower()) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """value of generate_pattern(10) is not correct"""
assert str(generate_pattern(10)) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """correct string value of generate_pattern(10) but incorrect case of letters"""

assert str(type(generate_pattern(6))) == """<class 'str'>""", """type of generate_pattern(6) is not str. generate_pattern(6) should be an str"""
assert str(len(generate_pattern(6))) == """33""", """length of generate_pattern(6) is not correct"""
assert str(generate_pattern(6).lower()) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """value of generate_pattern(6) is not correct"""
assert str(generate_pattern(6)) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """correct string value of generate_pattern(6) but incorrect case of letters"""

assert str(type(generate_pattern(3))) == """<class 'str'>""", """type of generate_pattern(3) is not str. generate_pattern(3) should be an str"""
assert str(len(generate_pattern(3))) == """12""", """length of generate_pattern(3) is not correct"""
assert str(generate_pattern(3).lower()) == """* *\n+++\n* *\n""", """value of generate_pattern(3) is not correct"""
assert str(generate_pattern(3)) == """* *\n+++\n* *\n""", """correct string value of generate_pattern(3) but incorrect case of letters"""

print('Success!')
```