## **Python String Exercises (Beginner to Advanced)**

### **Beginner Level: Basic String Manipulation**



#### **1. Concatenation**
Write a program that takes two strings as input and prints them concatenated together.



```python
s1 = input("Enter string 1")
s2 = input("Enter String 2")
print(s1 + s2)

```

Output:
```
Enter string 1 Har
Enter String 2 eni
Hareni
```

#### **2. String Indexing & Slicing**
- Print the first and last character of a given string.
- Print a substring containing only the first 5 characters.



```python
s = "Hello world!"
print(s[0:len(s)])
print(s[0],s[len(s)-1])
print(s[0:5])
```

Output:
```
Hello world!
H !
Hello
```

```python

```

#### **3. String Methods**
Convert a given string to:
- Uppercase
- Lowercase
- Title case (first letter capitalized)



```python
s = "HELLO"
up = ""
for i in s:
    if 65 <= ord(i) < 91:
        up += i
    else:
        num = ord(i)
        up += chr(num - 32)
print(up)
```

Output:
```
HELLO
```

```python
s = "hELlo"
lo = ""
for i in s:
    if 97 <= ord(i) < 123:
        lo += i
    else:
        num = ord(i)
        lo += chr(num + 32)
print(lo)
```

Output:
```
hello
```

```python
s = "helLo"
lo = ""
f = ""

for i in s:
    if 65 <= ord(i) < 91:
        i = chr(ord(i) + 32)
        lo += i
    else:
        lo += i
        
for x in lo:
    if x == lo[0]:
        f += chr(ord(lo[0]) - 32)
    else:
        f += x
    
  # 90 <= ord(lo[0]) < 123:
   #     f += chr(ord(lo[0]) + 32)
    #else:
     #   pass
       
print(f)
```

Output:
```
Hello
```

```python
97 + 26
```

Output:
```
123
```

```python
ord("h")
```

Output:
```
104
```

```python
90 + 32
```

Output:
```
122
```

```python
chr(122)
```

Output:
```
'z'
```

```python
ord("A")
```

Output:
```
65
```

#### **4. Find and Replace**
- Replace all occurrences of the word **"apple"** with **"orange"** in a given string.

---



```python
s = '''he is going to get an apple watch, 
at the apple store. 
The apple watch is called an iWatch.'''
w = "apple"
r = "orange"
check = "" #Storing the collected word into a temporary variable
d = 0

for i in s: # i = h
    l = ord(i) # l = 104
    while l != 32: # passes l != 32 # 
        check += i #check = "h" 
        d += 1 # d = 1
        l = 0 #l reset to 0
        n = s[d] #s[d] = n = "e"
        l += ord(n) # l = 0 + ord(e) l = 101 
        
print(check)

```

Output:
```
Cell In[60], line 10
    l = ord(i) l = 104
               ^
SyntaxError: invalid syntax
```

```python
s = '''he is going to get an apple watch, 
at the apple store. 
The apple watch is called an iWatch.'''
w = "apple"
r = "orange"
check = "" #Storing the collected word into a temporary variable
p = " " # p is space
new = "" # creating new string to add the sentence with the replaced word. 


for i in s:
    if i != p:
        check += i
    else:
        if check == w:
            new += r + " "
        else:
            new += check + " "
        #print(check,"-", new)
        check = ""
    
if check == w:
    new += r
else:
    new += check

print(new)
        
```

Output:
```
he is going to get an orange watch, 
at the orange store. 
The orange watch is called an iWatch.
```

```python
ord("e")
```

Output:
```
101
```

```python
ord("h")
```

Output:
```
104
```

```python
ord("e")
```

Output:
```
101
```

### **Intermediate Level: Loops with Strings**

#### **5. Iterate Over a String**
Write a program that prints each character of a string on a new line.



```python
s = "Hello"
for i in s:
    print(i)
```

Output:
```
H
e
l
l
o
```

```python

```

#### **6. Count Vowels in a String**
Write a function that counts the number of vowels (`a, e, i, o, u`) in a string.



```python
s = "I am a dog"
count = 0
vowels = "aeiouAEIOU"
for i in s:
    if i in vowels:
        count += 1
    else:
        pass
print(count)
```

Output:
```
4
```

```python
s = "I am a dog"
count = 0
vowels = "aeiou"
for i in s:
    if i in vowels:
        count += 1
    else:
        pass
print(count)
```

```python
ord("a")
```

Output:
```
97
```

```python
s = "i am a dog"
count = 0
vowels = "aeiou"
ns = ""
for i in s:
    i = (chr(ord(i) - 32))
    if i in vowels.upper():
        count += 1
    else:
        pass
print(ns)
print(count)
```

Output:
```
4
```

```python

```

```python
i = "h"
ns = (chr(ord(i) - 32))
print(ns)
```

Output:
```
H
```

```python
vowels = "aeiou"
print(vowels.upper())
```

Output:
```
AEIOU
```

#### **7. Reverse a String Using a Loop**
Write a program to reverse a string manually using a loop.



```python
s = "Come"
for i in range(len(s),0,-1):
    print(s[i])
```

Output:
```
IndexError                                Traceback (most recent call last)
Cell In[29], line 3
      1 s = "Come"
      2 for i in range(len(s),0,-1):
----> 3     print(s[i])

IndexError: string index out of range
```

```python
s = "Come"
for i in range(len(s)-1,-1,-1):
    print(s[i])
```

Output:
```
e
m
o
C
```

```python
s = "Come"
print(s[-(len(s))])
```

Output:
```
C
```

#### **8. Check if a String is a Palindrome**
Write a function to check if a given string is a palindrome (same forward and backward).



```python
s = "mNm"
s1 = ""
reverse = ""
for i in s:
    if 97 <= ord(i) < 122:
        i = chr(ord(i) - 32)
        s1 += i
    else:
        s1 += i

for x in range(len(s1)-1,-1,-1):
    reverse += s1[x]

if reverse == s1:
    print(f"{s} is a Pallindrome.")
else:
    print(f"{s} is not a Pallindrome.")
        
    
```

Output:
```
mNm is a Pallindrome.
```

```python
s = "Mam I am"
s1 = ""
reverse = ""
for i in s:
    if 97 <= ord(i) <= 122:
        i = chr(ord(i) - 32)
        s1 += i
    else:
        s1 += i

for x in range(len(s1)-1,-1,-1):
    reverse += s1[x]

if reverse == s1:
    print(f"{s} is a Pallindrome.")
else:
    print(f"{s} is not a Pallindrome.")
        
    
```

Output:
```
Mam I am is not a Pallindrome.
```


#### **9. Find the Longest Word in a Sentence**
Write a function that takes a sentence as input and returns the longest word.

---



```python
s  = "I was not a part of Schekjnolisjh"
check = ""
p = " "
longest_word = ""

for i in s:
    if i != p:
        check += i
    else:
        if len(check) > len(longest_word):
            longest_word = check
            
        elif len(check) == len(longest_word):
            longest_word1 = check

        check = ""

if len(check) > len(longest_word): # To check the last word
    longest_word = check
    print(f"Longest word in this sentence = {longest_word}")
else:
    print(f"Longest word in this sentence = {longest_word}")

    
```

Output:
```
Longest word in this sentence = Schekjnolisjh
```

```python
s  = "I was not a part of Red food."
```

```python

```

```python

```

```python
s  = "I was not a part of Germany's or Indianish food."
check = ""
p = " "
longest_word = ""

for i in s:
    if i != p:
        check += i
    else:
        if len(check) > len(longest_word):
            longest_word = check
            check = ""
        elif len(check) == len(longest_word):
            longest_word1 = check
        
        else:
            check = ""
        check = ""

if len(longest_word) > (len(check):
    print(f"Longest word in this sentence = {longest_word}")
elif len(longest_word) == len(longest_word1):
    print(f"Longest words in this sentence are = {longest_word} and {longest_word1}")

    
```

Output:
```
Cell In[135], line 20
    if len(longest_word) > (len(check):
                                      ^
SyntaxError: invalid syntax
```

```python
s  = "I was not a part of Schekjnolisjh"
check = ""
p = " "
longest_word = ""

for i in s:
    if i != p:
        check += i
    else:
        if len(check) > len(longest_word):
            longest_word = check
        check = ""

if len(check) > len(longest_word): # To check the last word
    longest_word = check
print(f"Longest word in this sentence = {longest_word}")
    
```

Output:
```
Longest word in this sentence = Schekjnolisjh
```

### **Advanced Level: Control Flow with Strings**

#### **10. Word Frequency Counter**
Count the occurrences of each word in a sentence and store the result in a dictionary.



```python
 s = "I am an apple, apple was a good girl"
check = ""
p = " "
count = 0
for i in s:
    if i != p:
        check += i
    else:
        for x in check:
            if check == s:
                count += 1
            
print(check)
```

Output:
```
Iamanapple,applewasagoodgirl
```

```python

```

#### **11. Extract Digits from a String**
Write a function that extracts only the digits from a given string.



```python
def dig(s):
    digs = []
    for i in s:
        if 48 <= ord(i) < 58: #ASCII character for digits 0 to 9
            digs += i
    print(digs)
    
s = "He3 5Dg"
dig(s)
```

Output:
```
['3', '5']
```

```python
s = "3"
type(s)
```

Output:
```
str
```

```python
ord("0")
```

Output:
```
48
```

```python
chr(90)
```

Output:
```
'Z'
```

```python
65 + 32
```

Output:
```
97
```

#### **12. Remove Special Characters from a String**
Write a function that removes special characters (e.g., `@, !, #, $, %`) from a string.



```python
def specialRemove(s):
    ns = ""
    for i in s:
        if (65 <= ord(i) <= 90) or (97 <= ord(i) <= 122) or (48 <= ord(i) <= 57) :
            ns += i
    return ns

t = "He§llo9$"
a = specialRemove(t)
print(a)
```

Output:
```
Hello9
```

```python

```

#### **13. Check for Anagrams**
Write a function that checks if two strings are anagrams (contain the same characters in a different order).



```python
def anagrams(s,t):
    letters = ""
    for i in s: # i = s
        if len(s) == len(t): # If both words are the same length
            for x in t: # x = a, x = r, x = t, x = s
                    if i == t: # False....... when x = s = True
                        letters += "" # letters = s
                    else: #move to the next
                        continue
        else:
            False


s = "star"
t = "arts"
print(anagrams(s,t))
```

Output:
```
None
```

```python
def anagrams(s,t):
    v = t
    if len(s) == len(t): # If both words are the same length
        for i in s: # i = s, i = t, 
            if i in t: # s in t = True , t in t = True 
                t = t.replace(i,"",1) # letters = $$$$
        
                
        if t == "" :
            print(f"{s} and {v} are Anagrams.")
            return True
        else:
            print(f"{s} and {v} are not Anagrams.")
            return False
            
    else:
        return False

    #if t == "":
     #   return True
    #else:
     #   return False


    

s = "star"
t = "arts"
print(anagrams(s,t))
```

Output:
```
star and arts are Anagrams.
True
```

```python
g = "Hallo"
print(g)
g = g.replace("l", "g", 1)  # Replaces the first occurrence of 'l' with an empty string
print(g)
```

Output:
```
Hallo
Haglo
```

Star and Arts

#### **14. Capitalize Every Alternate Letter in a String**
Write a function that capitalizes every alternate letter in a given string.



```python
def Cap(s):
    ns = ""
    for i in range(0,len(s),2):
        if i == 0:
            if 65 <= ord(s[i]) < 91:
                ns += s[i]
            else:
              ns += chr(ord(s[i]) - 32)
        else:
            for x in range(1,len(s)):
                if 65 <= ord(s[x]) < 91:
                    ns += s[x]
                else:
                    ns += chr(ord(s[i]) + 32)
              



s = "Hello"
print(Cap(s))

```

Output:
```
H
```

```python
def Cap(s):
    ns = ""
    cap = ""
    for i in s:
            if 97 <= ord(i) < 122:
                ns += i
            else:
              ns += chr(ord(i) + 32)
    #return ns
        
    for x in range(0,len(ns)): # Every alternate letter. 
        if x % 2 == 0:
            cap += chr(ord(ns[x]) - 32)
        else:
            cap += ns[x]
    
    return cap



s = "Hello"
print(Cap(s))

```

Output:
```
HeLlO
```

```python
def Cap(s):
    ns = ""
    for x in range(len(s)): # Every alternate letter. 
        if x % 2 == 0: # x = 0, x % 2 = 0
            if 97 <= ord(s[x]) < 123: #s[x] = s[0] = "h" if lowercase already
               ns += chr(ord(s[x]) - 32) # convert to upper case and add it to new string.
            else: #if it is uppercase
                ns += s[x] #leave as is and add it to new string
        else: #If x = 1
            if 65 <= ord(s[x]) < 91: # if uppercase
                ns += chr(ord(s[x]) + 32) 
            else:
                ns += s[x]
    
    return ns


s = "Hello"
print(Cap(s))

```

Output:
```
HeLlO
```

```python
def Cap(s):
    ns = ""
    for x in range(len(s)): # Every alternate letter. 
        if x % 2 == 0: # x = 0, x % 2 = 0
            if 97 <= ord(s[x]) < 123: #s[x] = s[0] = "h" if lowercase already
               ns += chr(ord(s[x]) - 32) # convert to upper case and add it to new string.
            else: #if it is uppercase
                ns += s[x] #leave as is and add it to new string
        else: #If x = 1
            if 65 <= ord(s[x]) < 91: # if uppercase
                ns += chr(ord(s[x]) + 32) 
            else:
                ns += s[x]
    
    return ns


s = "hEllo"
print(Cap(s))

```

Output:
```
HeLlO
```

```python
def Cap(s):
    ns = ""
    for x in range(0,len(s)): # Every alternate letter. 
        if x % 2 == 0: # x = 0, x % 2 = 0
            if 97 <= ord(s[x]) < 123: #s[x] = s[0] = "h"
               ns += chr(ord(s[x]) - 32)
            #else:
             #  pass 
        else:
            ns += s[x]
    
    return ns


s = "Hello"
print(Cap(s))

```

Output:
```
eLlO
```

```python
1 % 2
```

Output:
```
1
```

```python
s = "Hello"
ns = ""

for i in s:
        if 97 <= ord(i) < 122:
             ns += i
        else:
            ns += chr(ord(i) + 32)
print(ns)

cap = ""

for x in range(0,len(ns)):
    print(x,ns[x])
    if x % 2 == 0:
        cap += chr(ord(ns[x]) - 32)
    #cap += chr(ord(ns[x]) + 32)
    else:
        cap += ns[x]
print(cap)
```

Output:
```
hello
0 h
1 e
2 l
3 l
4 o
HeLlO
```

```python
ord("h")
```

Output:
```
104
```

```python
chr(90 + 32)
```

Output:
```
'z'
```

```python
s = "Hello"

    print(s[x])
```

Output:
```
e
l
l
o
```

```python

```

#### **15. Find the Most Common Character in a String**
Write a function that finds and returns the most frequently occurring character in a string.

---


```python
s = "Synonym"
maxChar = ""
maxCount = 0

for x in s: # x = h
    count = 0
    for y in s: # y = h (1st h)
        if y == x:
            count += 1
    if count > maxCount:
        maxCount = count
        maxChar = x
        
print(maxChar, maxCount)
            

```

Output:
```
y 2
```

```python
s = "hhlo"
count = 1
ts = s
iteration = 0
dicts = {}
for i in s: # i = h, s = hhllo, i = h, i = l
    iteration += 1
    s = s.replace(i,"",count) # s = 
    print("iteration:",iteration)
    print("s =", s)
    print("i =", i)
    for x in s: # s = "hllo", x = h
        print("x =",x)
        if x == i: #if h = h
            count += 1 # count = 2
            dicts[i] = count # {h : 2}
            print(dicts)
        else:
           dicts[i] = count
        
        
print(dicts)

```

Output:
```
iteration: 1
s = hlo
i = h
x = h
{'h': 2}
x = l
x = o
iteration: 2
s = lo
i = h
x = l
x = o
iteration: 3
s = o
i = l
x = o
iteration: 4
s = 
i = o
{'h': 2, 'l': 2}
```

```python
s = "hhlo"
count = 1
ts = s
iteration = 0
dicts = {}
for i in s: # i = h, s = hhllo, i = h, i = l
    if i in dicts:
       s = s.replace(i,"",count) 
    iteration += 1
     # s = 
    print("iteration:",iteration)
    print("s =", s)
    print("i =", i)
    for x in s: # s = "hllo", x = h
        print("x =",x)
        if x == i: #if h = h
            count += 1 # count = 2
            dicts[i] = count # {h : 2}
            print(dicts)
        else:
           dicts[i] = count
        
        
print(dicts)

```

Output:
```
iteration: 1
s = hhlo
i = h
x = h
{'h': 2}
x = h
{'h': 3}
x = l
x = o
iteration: 2
s = lo
i = h
x = l
x = o
iteration: 3
s = lo
i = l
x = l
{'h': 3, 'l': 4}
x = o
iteration: 4
s = lo
i = o
x = l
x = o
{'h': 3, 'l': 4, 'o': 5}
{'h': 3, 'l': 4, 'o': 5}
```

```python

```

```python
s = "hhllo"
count = 1
dicts = {}
for i in s: # i = h
    s = s.replace(i,"",1)
    for x in s: # s = "hllo", x = h
        if x == i: #if h = h
            count += 1 # count = 2
            dicts[i] = count
        else:
           dicts[i] = count
print(dicts)

```

Output:
```
{'h': 2, 'l': 3}
```

```python
s = "hhlo"
count = 1
iteration = 0
dicts = {}

# Outer loop to iterate through the string
for i in s:  # 'i' is the current character in the outer loop
    iteration += 1
    s_temp = s  # We use a temporary variable to preserve the original string

    # Replace the current character 'i' from the string
    s_temp = s_temp.replace(i, "", count)

    print("iteration:", iteration)
    print("s_temp =", s_temp)
    print("i =", i)
    
    # Initialize the count for this character
    count = 1

    # Loop through the modified string and count the occurrences of 'i'
    for x in s_temp:  # 'x' is each character in the modified string
        print("x =", x)
        if x == i:  # If 'x' matches 'i'
            count += 1  # Increment the count for 'i'
            dicts[i] = count  # Update the dictionary with the new count
            print("dicts =", dicts)
        else:
            dicts[i] = count  # Otherwise, store the current count for 'i'

# Print the final result
print("Final dicts:", dicts)

```

Output:
```
iteration: 1
s_temp = hlo
i = h
x = h
dicts = {'h': 2}
x = l
x = o
iteration: 2
s_temp = lo
i = h
x = l
x = o
iteration: 3
s_temp = hho
i = l
x = h
x = h
x = o
iteration: 4
s_temp = hhl
i = o
x = h
x = h
x = l
Final dicts: {'h': 1, 'l': 1, 'o': 1}
```

```python
s = "hhllo"

dicts = {}
for i in s: # i = h
    s = s.replace(i,"",1)
    count = 1
    if i in s: # s = "hllo", x = h #if h = h
        count += 1 # count = 2
        dicts[i] = count
    else:
       dicts[i] = count

        
print(dicts)

```

Output:
```
{'h': 1, 'l': 1, 'o': 1}
```

```python
s = "hhllo"
dicts = {}

for i in s:  # iterate through each character in the string
    count = 1  # initialize count to 1 for the first occurrence of the character
    s = s.replace(i, "", 1)  # remove the first occurrence of 'i'
    
    if i in s:  # if 'i' appears again in the string
        count += 1  # increment the count
        dicts[i] = count  # store the count in the dictionary
    else:
        dicts[i] = count  # store the count for the last occurrence

print(dicts)

```

Output:
```
{'h': 1, 'l': 1, 'o': 1}
```

```python
s = "hhllo"

dicts = {}
original_s = s  # Keep the original string intact

for i in original_s:  # Iterate through each character in the string
    if i not in dicts:  # Check if the character has already been counted
        count = 1  # Initialize count for the character

        s = original_s  # Reset the string to its original value
        
        # Loop through the remaining occurrences of the character
        for x in s:
            if x == i:  # Check if the character matches
                count += 1  # Increment count
        dicts[i] = count  # Store the count in the dictionary

print(dicts)

```

Output:
```
{'h': 3, 'l': 3, 'o': 2}
```

```python
s = "hhllo"

dicts = {}

# Iterate through each character in the string
for i in s:  
    if i not in dicts:  # Only process the character if it's not in the dictionary already
        count = 1  # Start the count for this character

        # Make a copy of the string to avoid modifying the original string in the loop
        remaining_string = s.replace(i, "", 1)  # Remove the first occurrence of 'i'
        
        # Now count the remaining occurrences of 'i' in the string
        while i in remaining_string:
            count += 1
            remaining_string = remaining_string.replace(i, "", 1)  # Remove one more occurrence of 'i'
        
        dicts[i] = count  # Store the final count in the dictionary

print(dicts)

```

Output:
```
{'h': 2, 'l': 2, 'o': 1}
```

#### **16. Slicing**

Copy the provided `text` to the returned variable `new_string` after replacing each vowel with a dash (`-`) using an `if-else` statement with a **single condition**, that consists of a single `in` operator.

```python
text = "Hello"
new_string = ""
vowel = "aeiouAEIOU"
for i in text:
    if i in vowel:
        new_string += "-"
    else:
        new_string += i
print(new_string)
```

Output:
```
H-ll-
```

```python

```

```python

```

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each vowel with a dash (`-`) using an `if-else` statement with a **single condition**, that consists of a single `in` operator.

<!-- grade_id: cell-38241276a7919b5a, points: 0, grade: false, solution: true -->
```python
def vowels(text):
  
    new_string = ""
    vowel = "aeiouAEIOU"
    for i in text:
        if i in vowel:
            new_string += "-"
        else:
            new_string += i
    return new_string

text = "Hello World"

print(vowels(text))
```

Output:
```
H-ll- W-rld
```

<!-- grade_id: cell-f2db875188f9b58e, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(vowels("Hello World"))) == """<class 'str'>""", """type of vowels(\"Hello World\") is not str. vowels(\"Hello World\") should be an str"""
assert str(len(vowels("Hello World"))) == """11""", """length of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World").lower()) == """h-ll- w-rld""", """value of vowels(\"Hello World\") is not correct"""
assert str(vowels("Hello World")) == """H-ll- W-rld""", """correct string value of vowels(\"Hello World\") but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

Copy the provided `text` to the returned variable `new_string` after replacing each **letter appearing in the given string** `replace` with a dash. Note: you should replace the text given in the variable `text`, *not* the text in `replace`.


*Warning:* Your program should work **regardless** which letters are given in the `replace` string, and even if no letters are provided at all (when `replace` is the empty string).

<!-- grade_id: cell-dba35b06c5e38d24, points: 0, grade: false, solution: true -->
```python
def general_replace(text, replace):
    new_string = ""
    for i in text:
        if i in replace:
            new_string += "-"
        else:
            new_string += i
    return new_string

text = "Hello World"
replace = ""
print(general_replace(text, replace))
```

Output:
```
Hello World
```

<!-- grade_id: cell-ea7b28540a173711, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(general_replace("Hello World", "aeiou"))) == """<class 'str'>""", """type of general_replace(\"Hello World\", \"aeiou\") is not str. general_replace(\"Hello World\", \"aeiou\") should be an str"""
assert str(len(general_replace("Hello World", "aeiou"))) == """11""", """length of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou").lower()) == """h-ll- w-rld""", """value of general_replace(\"Hello World\", \"aeiou\") is not correct"""
assert str(general_replace("Hello World", "aeiou")) == """H-ll- W-rld""", """correct string value of general_replace(\"Hello World\", \"aeiou\") but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (1P.)

For the given string `text` return a string that consists of `n`-times the original string, for example for "a" and 5 you should return "aaaaa".

<!-- grade_id: cell-27137fae67d5c4b5, points: 0, grade: false, solution: true -->
```python
def repeat(text, n):
    return (text * n)

print(repeat("a", 5))
```

Output:
```
aaaaa
```

<!-- grade_id: cell-8defb15d0b45d3f9, points: 1, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(repeat("a", 5))) == """<class 'str'>""", """type of repeat(\"a\", 5) is not str. repeat(\"a\", 5) should be an str"""
assert str(len(repeat("a", 5))) == """5""", """length of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5).lower()) == """aaaaa""", """value of repeat(\"a\", 5) is not correct"""
assert str(repeat("a", 5)) == """aaaaa""", """correct string value of repeat(\"a\", 5) but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

For the given input `text`, check whether the string in `text` is the same as its reverse (the string read backwards). For example, your function should return `True` for "level" but not for "lever".

*Hint: You can use slice notation for the reversing.*

<!-- grade_id: cell-d848906595b1a870, points: 0, grade: false, solution: true -->
```python
def backwards_equality(text):
    reverse = ""
    for i in range(len(text)-1,-1,-1):
        reverse += text[i]
    if reverse == text:
        return True
    else:
        return False







text = "level"
print(backwards_equality(text))
```

Output:
```
True
```

<!-- grade_id: cell-20b294b518a836b2, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(backwards_equality("level"))) == """<class 'bool'>""", """type of backwards_equality(\"level\") is not bool. backwards_equality(\"level\") should be a bool"""
assert str(backwards_equality("level")) == """True""", """boolean value of backwards_equality(\"level\") is not correct"""

assert str(type(backwards_equality("lever"))) == """<class 'bool'>""", """type of backwards_equality(\"lever\") is not bool. backwards_equality(\"lever\") should be a bool"""
assert str(backwards_equality("lever")) == """False""", """boolean value of backwards_equality(\"lever\") is not correct"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

In the return variable `new_string`, return a string where the substring of the given string `s` between the indices `start` and `end` (inclusive) is reversed, while the rest of the string remains unchanged.  

**Make sure that your function returns a reasonable result for every input (assume `start` and `end` are valid indices, and `start <= end`).**

*Hint: Use string slicing and indexing.*

```python
s = "Hello"
start = 1
end = 2
print(s[0:len(s)])
reverse = s[start:end + 1]
print(reverse)
```

Output:
```
Hello
el
```

```python
s = "Hello"
new_string = ""
for x in range(len(s)):
    rs = s[start:end + 1]
print(rs)
nrs = ""
for y in range(len(rs)-1,-1,-1):
    nrs += rs[y]
    print(nrs)
```

Output:
```
el
l
le
```

```python
s = "Hello"
new_string = ""
start = 1
end = 2
print(s[0])
```

```python
s = "Hello"
new_string = ""
start = 1
end = 2
for x in range(len(s)):
    rs = s[start:end + 1]
print(rs)
nrs = ""
for y in range(len(rs)-1,-1,-1):
    nrs += rs[y]
    print(nrs)
```

<!-- grade_id: cell-70c2a1b68e7165d8, points: 0, grade: false, solution: true -->
```python
def reverse_substring(s, start, end):
    new_string = ""
    for x in range(len(s)):
        if x == start:
            for y in range(end, start - 1, -1): # end = 2, start = 
                new_string += s[y]
        else:
            new_string += s[x]
    return new_string

s = "Hello"
print(reverse_substring(s,1,2))
   
```

Output:
```
Hlello
```

```python
s = "Hello"
for i in range(len(s)):
    print(-i, s[-i])
print(s[-1])
```

Output:
```
0 H
-1 o
-2 l
-3 l
-4 e
o
```

```python
start = 1
end = 4
s = "Hello"
print(s[::-1])
```

Output:
```
olleH
```

```python
def reverse_substring(s, start, end):
    new_string = s[:start] + s[start:end+1][::-1] + s[end+1:]
    return new_string      #   2  : 0 : -1  
    
s = "Hello"
    
print(reverse_substring(s,1,2))
   
```

Output:
```
Hlelo
```

<!-- grade_id: cell-c793f233248fc05c, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(reverse_substring("Hello World!", 3, 6))) == """<class 'str'>""", """type of reverse_substring(\"Hello World!\", 3, 6) is not str. reverse_substring(\"Hello World!\", 3, 6) should be an str"""
assert str(len(reverse_substring("Hello World!", 3, 6))) == """12""", """length of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6).lower()) == """helw olorld!""", """value of reverse_substring(\"Hello World!\", 3, 6) is not correct"""
assert str(reverse_substring("Hello World!", 3, 6)) == """HelW olorld!""", """correct string value of reverse_substring(\"Hello World!\", 3, 6) but incorrect case of letters"""

assert str(type(reverse_substring("Hi", 0, 1))) == """<class 'str'>""", """type of reverse_substring(\"Hi\", 0, 1) is not str. reverse_substring(\"Hi\", 0, 1) should be an str"""
assert str(len(reverse_substring("Hi", 0, 1))) == """2""", """length of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1).lower()) == """ih""", """value of reverse_substring(\"Hi\", 0, 1) is not correct"""
assert str(reverse_substring("Hi", 0, 1)) == """iH""", """correct string value of reverse_substring(\"Hi\", 0, 1) but incorrect case of letters"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

Write a function, that given two strings `text1` and `text2`, checks whether they are equal if case is ignored, so "hELLO" should be equal to "HellO".

*Remember that you are not allowed to use build-in string methods such as `s.isupper()`.*

*Hint: Use lexicographic comparison.*

*Hint: Remember `break` and `continue`.*

```python
g = "Hello"
g = g.replace(g[1],g[-1])
print(g)
```

Output:
```
Hollo
```

<!-- grade_id: cell-8a5009af8b821f0f, points: 0, grade: false, solution: true -->
```python
def equal_without_case(text1, text2):
    for i in text1:
        if 65 <= ord(i) < 91:
            new = chr(ord(i) + 32)
            text1 = text1.replace(i,new)
        #else:
         #   continue
    for x in text2:
        if 65 <= ord(x) < 91:
            n = chr(ord(x) + 32)
            text2 = text2.replace(x,n)

    if text1 == text2:
        return True
    else:
        return False
        


t1 = "HELLO"
t2 = "HeLlO"
equal_without_case(t1,t2)
```

Output:
```
True
```

##### Without using in built methods like  `s.replace()`

```python
def equal_without_case(text1, text2):
    lt1 = ""
    for i in text1:
        if 65 <= ord(i) < 91: #If letter is capital
            lt1 += chr(ord(i) + 32) # Convert to small and add it to lt
        else:
            lt1 += i
            
    lt2 = ""  
    for x in text2: #If letter is capital
        
        if 65 <= ord(x) < 91:
            lt2 += chr(ord(x) + 32)
        else:
            lt2 += x

    if lt1 == lt2:
        return True
    else:
        return False
        


t1 = "HELLO"
t2 = "HeLlO"
equal_without_case(t1,t2)
```

Output:
```
True
```

##### Using break and continue

```python
def equal_without_case(text1, text2):
    lt1 = ""
    for i in text1:
        if 65 <= ord(i) < 91: #If letter is capital
            lt1 += chr(ord(i) + 32) # Convert to small and add it to lt
            continue
        lt1 += i
            
    lt2 = ""  
    for x in text2: #If letter is capital
        
        if 65 <= ord(x) < 91:
            lt2 += chr(ord(x) + 32)
            continue
        lt2 += x

    if lt1 == lt2:
        return True
    else:
        return False
        


t1 = "HEwd"
t2 = "HeLlO"
equal_without_case(t1,t2)
```

Output:
```
False
```

<!-- grade_id: cell-81dbc76ec9ba6ce1, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(equal_without_case("hELLO", "HellO"))) == """<class 'bool'>""", """type of equal_without_case(\"hELLO\", \"HellO\") is not bool. equal_without_case(\"hELLO\", \"HellO\") should be a bool"""
assert str(equal_without_case("hELLO", "HellO")) == """True""", """boolean value of equal_without_case(\"hELLO\", \"HellO\") is not correct"""

assert str(type(equal_without_case("hello", "HELLO"))) == """<class 'bool'>""", """type of equal_without_case(\"hello\", \"HELLO\") is not bool. equal_without_case(\"hello\", \"HELLO\") should be a bool"""
assert str(equal_without_case("hello", "HELLO")) == """True""", """boolean value of equal_without_case(\"hello\", \"HELLO\") is not correct"""

print('Success!')
```

Output:
```
Success!
```

## Task (2P.)

In this task, you are meant to re-create simple [ASCII art](https://en.wikipedia.org/wiki/ASCII_art).  
Write a function that returns the following pattern for different given heights:

### Example Height = 10

```
*        *
 *      *
  *    *
   *  *
++++++++++
   *  *
  *    *
 *      *
*        *
```

### Example Height = 7

```
*     *
 *   *
  * *
+++++++
  * *
 *   *
*     *
```

The function should receive the height as an input and return a string containing the pattern.

**Hints**:
1. To add a new line to a string `s`, use `s += "\n"`. To see the resulting string formatted correctly, you must use `print(s)`; evaluating `s` alone won’t display the line breaks correctly in JupyterHub.
2. Treat even and odd heights differently to maintain symmetry, as seen in the examples.
3. The lowest height for which the function must produce a meaningful string is `3`.

```python
height = 7
for i in range(height-1,0,-2):
    for x in range(1,i):
        print("$", end = "")
    print()
    
```

Output:
```
$$$$$
$$$
$
```

```python
for x in range(1,(height // 2) + 1):
    for y in range(x):
        print("!", end = "")
    print("*")
print()
```

Output:
```
!*
!!*
!!!*
```

```python
for x in range(1,(height // 2) + 1):
    for y in range(x):
        print("!", end = "")
    print("*")  

for i in range(height-2,0,-2):
       for j in range(1,i):
           print("$", end = "")
       print()
    

```

Output:
```
!*
!!*
!!!*
$$$$
$$
```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print("*")  

for i in range(y):
       for j in range(5,0,-2):
           print("$", end = "")
       print()
    

```

Output:
```
*
!*
!!*
$$$
```

```python
for i in range(5,0,-2):
    print(i)
```

Output:
```
5
3
1
```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print("*")  

for i in range(height-1,0,-2):
       for j in range(1,i):
           print("$", end = "")
       print()
    

```

Output:
```
*
!*
!!*
$$$$$
$$$
$
```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print("*")  

for i in range(height-1,0,-2):
       for j in range(1,i):
           print("$", end = "")
       print()
    

```

Output:
```
*
!*
!!*
$$$$$
$$$
$
```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print()  

for i in range(height-1,0,-2):
    print("*",end = "")
    for j in range(1,i):
       print("$", end = "")
    print()
print()
    

```

Output:
```
!
!!
*$$$$$
*$$$
*$
```

```python
star = "$"
for i in range(1,5):
    for j in range(i):
        print(i,"*",end = "")
    print()
print(star)
    
```

Output:
```
1 *
2 *2 *
3 *3 *3 *
4 *4 *4 *4 *
$
```

```python

```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print("*")  

    for i in range(x + (height-1),0,-2):
           print("$", end = "")
print()
    

```

Output:
```
*
$$$!*
$$$$!!*
$$$$
```

```python
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
        
    print("*")  

    for j in range(1, height - (2 * x)):  # Keeping your logic for 'j' loop
        print("$", end="")
    
    print()
    

```

Output:
```
*
$$$$$$
!*
$$$$
!!*
$$
```

```python
height = 7
for x in range((height // 2)):
    for y in range(x):
        print("!", end = "")
    print("*")  
```

Output:
```
*
!*
!!*
```

```
*     *    x = 1, $ = 7 - 2 = 5
 *   *     $ = 3
  * *      $ = 1
+++++++
  * *
 *   *
*     *

```python
height = 7
half = height // 2
for x in range(1, half + 1):
    for y in range(x):
        space = "$" * (height - x)
```

```python
height = 7
half = height // 2
for x in range(1, half + 1):
    for y in range(x):
        space = "$" * (height - x)
    for j in range(x, x + 2):
        print("*", space, end = "")
    print()

    #for z in range(1, half + 1):
       # if z == x:
            
    
for i in range(half,half + 1):
    for j in range(height):
        print("+", end = "")
    print()
```

Output:
```
* $$$$$$* $$$$$$
* $$$$$* $$$$$
* $$$$* $$$$
+++++++
```

```python
height = 7

# Printing '!' part
for x in range(height // 2):
    for y in range(x):
        print(" ", end="")  # Print '!' without moving to a new line
    
    print("*", end="")  # Print '*' in the same line as '!'

    # Start the second loop for '$' (your logic)
    for j in range(1, height - (2 * x)):  # Keeping your logic for 'j' loop
        print("$", end="")

    print("*", end="")
    print()

half = height // 2
for i in range(half,half + 1):
    for j in range(height):
        print("+", end = "")
      # Move to the next line
```

Output:
```
*$$$$$$*
 *$$$$*
  *$$*
+++++++
```

```python
for j in range(1, 7 - (2 * x)):  # Keeping your logic for 'j' loop
        print("$", end="")

   
print()
```

Output:
```
$$
```

```python
 for j in range(1, height - (2 * x)):  # Keeping your logic for 'j' loop
        print("$", end="")
```

Output:
```
$$
```

```python
height = 7

# Printing the upper triangle part
for x in range(height // 2):
    for y in range(x):  
        print(" ", end="")  # Print leading spaces
    
    print("*", end="")  # Print the first '*'
    
    # Printing spaces between '*'
    for j in range(1, height - (2 * x)):  # Keeping your logic for spacing
        print(" ", end="")

    print("*")  # Print the second '*' and move to the next line

# Printing the base row of '+'
half = height // 2
for i in range(half, half + 1):
    for j in range(height):
        print("+", end = "")  # Print '+' without space
    print()  # Move to the next line

height = 7
half = height // 2 # Adjusted for the given pattern

# Printing the upper triangle part
for x in range(half // 2):
    for y in range(height - x - 1):  
        print(" ", end="")  # Print leading spaces for alignment
    
    print("*", end="")  # Print the first '*'
    
    # Printing spaces between '*' characters
    if x > 0:  # Ensure spaces appear only when x > 0
        for j in range(2 * x - 1):  
            print(" ", end="")
        print("*")  # Print the second '*'
    else:
        print()  # Move to the next line

```

Output:
```
*      *
 *    *
  *  *
+++++++
      *
```

```python
height = 7
half = height // 2
for l in range(half + 1, height):
    for m in range(height-l):  
        print("$", end="")  # Print leading spaces 
    print()  #Print the first'*'
    

```

Output:
```
$$$
$$
$
```

```python
for l in range(4, 7):
    for m in range(7-l):  
        print("$", end="")
    print()# Print leading spaces 
```

Output:
```
$$$
$$
$
```

```python
8 // 2
```

Output:
```
4
```

```python
for x in range(7):
    for y in range(7-x):
        print("$", end = "")
    print()
```

Output:
```
$$$$$$$
$$$$$$
$$$$$
$$$$
$$$
$$
$
```

```python
 # Printing spaces between '*'
    for j in range(1, height - (2 * x)):  # Keeping your logic for spacing
        print(" ", end="")

    print("*")
```

```python
height = 7
for y in range(1, (height // 2) + 1):
        print("$", end = "")
    #print()
```

Output:
```
$$$
```

```python
height = 7
for x in range(1,(height // 2) + 1):
    for y in range(x):
        print("$", end = "")
    print()
```

Output:
```
$
$$
$$$
```

```python
7 // 2
```

Output:
```
3
```

```python
height = 7
for x in range(1, height + 1):
    for y in range(0,(x + 1) // 2):
        print("$", end = "")
    print()
```

Output:
```
$
$
$$
$$
$$$
$$$
$$$$
```

<!-- grade_id: cell-2d0414bef2c32926, points: 0, grade: false, solution: true -->
```python
def generate_pattern(height):
    pattern = ""
# YOUR SOLUTION HERE
    return pattern
```

<!-- grade_id: cell-1a8eb14a86a88b76, points: 2, grade: true, solution: false -->
```python
# This tests your implementation:
from hashlib import sha1
assert str(type(generate_pattern(10))) == """<class 'str'>""", """type of generate_pattern(10) is not str. generate_pattern(10) should be an str"""
assert str(len(generate_pattern(10))) == """87""", """length of generate_pattern(10) is not correct"""
assert str(generate_pattern(10).lower()) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """value of generate_pattern(10) is not correct"""
assert str(generate_pattern(10)) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """correct string value of generate_pattern(10) but incorrect case of letters"""

assert str(type(generate_pattern(6))) == """<class 'str'>""", """type of generate_pattern(6) is not str. generate_pattern(6) should be an str"""
assert str(len(generate_pattern(6))) == """33""", """length of generate_pattern(6) is not correct"""
assert str(generate_pattern(6).lower()) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """value of generate_pattern(6) is not correct"""
assert str(generate_pattern(6)) == """*    *\n *  *\n++++++\n *  *\n*    *\n""", """correct string value of generate_pattern(6) but incorrect case of letters"""

assert str(type(generate_pattern(3))) == """<class 'str'>""", """type of generate_pattern(3) is not str. generate_pattern(3) should be an str"""
assert str(len(generate_pattern(3))) == """12""", """length of generate_pattern(3) is not correct"""
assert str(generate_pattern(3).lower()) == """* *\n+++\n* *\n""", """value of generate_pattern(3) is not correct"""
assert str(generate_pattern(3)) == """* *\n+++\n* *\n""", """correct string value of generate_pattern(3) but incorrect case of letters"""

print('Success!')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 3
      1 # This tests your implementation:
      2 from hashlib import sha1
----> 3 assert str(type(generate_pattern(10))) == """<class 'str'>""", """type of generate_pattern(10) is not str. generate_pattern(10) should be an str"""
      4 assert str(len(generate_pattern(10))) == """87""", """length of generate_pattern(10) is not correct"""
      5 assert str(generate_pattern(10).lower()) == """*        *\n *      *\n  *    *\n   *  *\n++++++++++\n   *  *\n  *    *\n *      *\n*        *\n""", """value of generate_pattern(10) is not correct"""

NameError: name 'generate_pattern' is not defined
```

```python

```